```python
# Initialize Otter
import otter
grader = otter.Notebook("DSC100_Lab6.ipynb")
```

# Lab 06: Visualization, Transformations, and KDEs

In this lab, you will get some practice plotting, applying data transformations, and working with Kernel Density Estimators (KDEs). 

In the first part of the lab, we will be working with data from the World Bank containing various statistics for countries and territories around the world. In the second part of the lab, we will work with a real-world dataset about bike-sharing. 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others, please **include their names** below.

**Collaborators:** *list collaborators here*

## 🏎️ Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

Finally, unless we state otherwise, **do not use for loops or list comprehensions**. The majority of this assignment can be done using built-in commands in `pandas` and `NumPy`.  Our autograder isn't smart enough to check, but you're depriving yourself of key learning objectives if you write loops / comprehensions, and you also won't be ready for the midterm next week.


## ✅ Grading
Grading is broken down into auto-graded answers and free responses. For auto-graded answers, the results of your code are compared to provided and/or hidden tests. For free response, readers will evaluate how well you answered the question and/or fulfilled the requirements of the question. Most of the plots will be manually graded. Your plots should be *similar* to the given examples. We will accept small variations, such as color differences or slight variations in scale. However, it is in your best interest to make the plots as similar as possible, as similarity is subject to the readers. 

**Note that for ALL plotting questions from here on out, we will expect descriptive titles, axis labels, legends, etc. The following question serves as a good guideline on what is "enough": If I directly downloaded the plot and viewed it, would I be able to tell what was being visualized without knowing the question?** 

### Score breakdown



Question | Manual | Points
--- |---| ---
1a |Yes | 1
1b |Yes | 1
1c |Yes | 1
1di |Yes | 1
1dii|Yes | 1
2a |Yes | 1
2b |No | 1
2c |Yes | 1
3a |Yes | 1
3b |Yes | 1
3c |Yes | 1
4a |Yes |2
5a |No |3
5b |No |3
6a |Yes |2
6b |Yes |1
6c |Yes |2
6d |Yes |2
7a |Yes |3
7b |Yes |1
7c |Yes |3
8a |No |1
8b |Yes |2
8c |Yes |2
8d |Yes |1
9a |No |1
9b |Yes |4
9c |Yes |2
10a |Yes |2
10b |Yes |2
Total |29 |50

```python
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from IPython.display import display, Markdown, Latex

plt.style.use('fivethirtyeight') # Use plt.style.available to see more styles
sns.set()
sns.set_context("talk")
%matplotlib inline
```

<br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Loading Data

Let's load some World Bank data into a `pd.DataFrame` object named ```wb```.

```python
wb = pd.read_csv("data/world_bank_misc.csv", index_col=0)
wb.head()
```

This table contains some interesting columns. Run this cell to take a look:

```python
display(Markdown('\n'.join(fr'{i + 1}. {c}'.replace('$', r'\$') for i, c in enumerate(wb.columns))))
```

<hr style="border: 1px solid #fdb515;" />

## Part 1: Scaling
In the first part of this lab, we will look at the distribution of values for the *combined adult literacy rate* as well as *the gross national income per capita*. The code below creates a copy of the `DataFrame` `wb` that contains only the two `Series` we want and then drops all rows that contain `null` values in either column.

**Note:** For this lab, we are dropping `null` values without investigating them further. However, this is generally not the best practice and can severely affect our analyses.

Here, the combined literacy rate is the sum of the female and male literacy rates as reported by the World Bank. 0 represents no literacy, and 200 represents total literacy by both genders that are included in the World Bank's dataset.

In this lab, we will be using the `sns.histplot`, `sns.rugplot`, and `sns.displot` functions to visualize distributions. You may find it useful to consult the `seaborn` documentation on [distributions](https://seaborn.pydata.org/tutorial/distributions.html) and [functions](https://seaborn.pydata.org/tutorial/function_overview.html) for more details.

```python
# Run this next line to create a DataFrame with the appropriate index; no further action is needed.
df = pd.DataFrame(index=wb.index)

# Run this next line to create copies of the Series we want; no further action is needed.
# The "\" syntax lets us continue to write on the next line.
df['lit'] = wb['Adult literacy rate: Female: % ages 15 and older: 2005-14'] \
            + wb["Adult literacy rate: Male: % ages 15 and older: 2005-14"]
df['inc'] = wb['Gross national income per capita, Atlas method: $: 2016']

# Run this next line to drop all records that have a NaN value in either column; no further action is needed.
df.dropna(inplace=True)
print("Original records:", len(wb))
print("Final records:", len(df))
```

```python
df.head(5)
```

<br>

---
### Question 1a

Suppose we wanted to build a histogram of our data to understand the distribution of literacy rates and income per capita individually. We can use [`countplot`](https://seaborn.pydata.org/generated/seaborn.countplot.html) in `seaborn` to create bar charts from categorical data. 

```python
sns.countplot(data=df, x="lit")
plt.xlabel("Combined literacy rate (ages >=15; 2005-2014)")
plt.title('World Bank Combined Adult Literacy Rate');
```

```python
sns.countplot(data=df, x="inc")
plt.xlabel('Gross national income per capita, Atlas method: $: 2016', fontsize=15)
plt.title('World Bank Gross National Income Per Capita');
```

<!-- BEGIN QUESTION -->

In the cell below, explain why `countplot` is NOT the right tool for visualizing the distribution of our data.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

---
### Question 1b

In the cell below, create a plot of **income per capita** (the second plot above) using the [`histplot`](https://seaborn.pydata.org/generated/seaborn.histplot.html) function. 
Don't forget to title the plot and label axes!

Optional: Create two side-by-side subplots, where the left subplot is **income**, and the right subplot is **literacy**. 

**Hint:** Reference the code from `q1a` to get started.

<!-- grade_id: q2a-answer, points: 1, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

You should see histograms that show the counts of how many data points appear in each bin. `displot` uses a heuristic called the Freedman-Diaconis rule [(original paper)](https://link.springer.com/content/pdf/10.1007/bf01025868) to automatically identify the best bin sizes. Although it is possible to set the bins yourself, we will not be doing so in this lab.

In the cell below, we explore overlaying a `rugplot` on top of a `displot` using `rugplot`. Note that the `rugplot` is hard to see.

<!-- grade_id: cell-72d2c195130debd8, points: 1, grade: true, solution: true -->
```python
# Run this cell to overlay a rugplot on the displot; no further action is needed.
sns.displot(data=df, x="inc")
sns.rugplot(data=df, x="inc")
plt.xlabel('Gross national income per capita, Atlas method: $: 2016', fontsize=15)
plt.title('World Bank Gross National Income Per Capita');
```

In the cell below, we explore overlaying a `rugplot` on top of a `histplot`. One way to make it easier to see the difference between the rug plot and the bars is to set a different color. For example:

```python
# Run this cell to overlay a rugplot on the histplot; no further action is needed.
sns.histplot(data=df, x="inc", color="lightsteelblue")
sns.rugplot(data=df, x="inc")
plt.xlabel('Gross national income per capita, Atlas method: $: 2016', fontsize=15)
plt.title('World Bank Gross National Income Per Capita');
```

There is also another function called `kdeplot`, which plots a Kernel Density Estimate (KDE) as described in class and will be covered in more detail later in this lab. 

Rather than manually calling `histplot`, `rugplot`, and `kdeplot` to plot histograms, rug plots, and KDE plots, respectively, we can instead modify how we use `displot` to simultaneously generate all three types of plots and adjust colors automatically for visibility. Using the [documentation](https://seaborn.pydata.org/generated/seaborn.displot.html) for `displot`, make a plot of the income data that includes a histogram, rug plot, and KDE plot. 

<!-- grade_id: cell-1aaac759f52a2269, points: 1, grade: true, solution: true -->
```python
# Run this cell to generate a displot; no further action is needed.
sns.displot(data=df, x="inc", rug=True, kde=True)
plt.xlabel('Gross national income per capita, Atlas method: $: 2016', fontsize=15)
plt.title('World Bank Gross National Income Per Capita');
```

You should see roughly the same histogram as before. However, now you should see an overlaid smooth line. This is the kernel density estimate discussed in class. 

Above, the y-axis is labeled by the counts. We can also label the y-axis by the density. An example is given below, this time using the literacy data from the beginning of this lab.

```python
# Run this cell to generate a displot; no further action is needed.
sns.displot(data=df, x="lit", rug=True, kde=True, stat="density")
plt.xlabel("Combined adult literacy rate (%) ages >=15; 2005-14", fontsize=14)
plt.title('World Bank Combined Adult Literacy Rate');
```

Observations:

* You'll see that the y-axis value is no longer the count. Instead, it is a value such that the total **area** in the histogram is 1. For example, the area of the last bar is approximately $18.67 \times 0.03 = 0.56$.
* The KDE is a smooth estimate of the distribution of the given variable. The area under the KDE is also 1. While it is not obvious from the figure, some of the area under the KDE is beyond 100% literacy. In other words, the KDE is non-zero for values greater than 100%. This, of course, makes no logical sense. Nonetheless, it is a mathematical feature of the KDE.

We'll talk more about KDEs later in this lab.

<!-- BEGIN QUESTION -->

<br>

---
### Question 1c

Looking at the income data, it is difficult to see the distribution among low-income countries because they are all scrunched up at the left side of the plot. The KDE also has a problem where the density function has a lot of area below 0. 

Transforming the `inc` data logarithmically gives us a more symmetric distribution of values. This can make it easier to see patterns.

In addition, summary statistics like the mean and standard deviation (square root of the variance) are more stable with symmetric distributions.

In the cell below, make a distribution plot of `inc` with the data transformed using `np.log10`. Set the arguments `kde=True` and `stat='density'`. 

Don't forget to title the plot and label axes! If your title(s) are too long and get cut off from the image, you can optionally set the `fontsize=` argument to adjust the font size. Take a look at the code in `q1b` for examples!

**Hint:** Seaborn's plotting functions like `displot` ([documentation](https://seaborn.pydata.org/generated/seaborn.displot.html)) have quite flexible parameters. Instead of setting `x=` as a column name, you can pass in a `Series` of the data you'd like to visualize.

<!-- grade_id: q2c-answer, points: 1, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

When a distribution has a long right tail (right-skewed distributions), a log transformation often does a good job of symmetrizing the distribution, as it did here. Long right tails are common with variables that have a lower limit on the values. 

On the other hand, long left tails are common with distributions of variables that have an upper limit, such as percentages (can't be higher than 100%) and GPAs (can't be higher than 4). That is the case for the literacy rate. Typically taking a power transformation such as squaring or cubing the values can help symmetrize a left-skewed distribution.

In the cell below, we will make a distribution plot of `lit` with the data transformed using a power, i.e., raise `lit` to the 2nd, 3rd, and 4th power. We plot the transformation with the 4th power below.


```python
# Run this cell to generate the plot; no further action is needed.
ax = sns.displot(data=df, x=df['lit']**4, kde=True, stat='density')
plt.xlabel("Combined adult literacy rate (%) ages >=15; 2005-2014", fontsize=14)
plt.ylabel('Density')
plt.title('World Bank Combined Adult Literacy Rate (4th power)', pad=30);
```

<!-- BEGIN QUESTION -->

<br>

---
### Question 1d

If we want to examine the relationship between the combined adult literacy rate and the gross national income per capita, we need to make a scatter plot. 

In the cell below, create a scatter plot of untransformed income per capita and literacy rate using the `sns.scatterplot` function. Make sure to label both axes using `plt.xlabel` and `plt.ylabel`. Don't forget to title the plot and label axes! If your title(s) are too long and get cut off from the image, you can optionally set the `fontsize=` argument to adjust the font size. Take a look at the code above for examples!

```python
...
```

<!-- END QUESTION -->

We can better assess the relationship between two variables when they have been straightened because it is easier for us to recognize linearity.

In the cell below, we see a scatter plot of log-transformed income per capita against literacy rate.

```python
sns.scatterplot(data=df, x='lit', y=np.log10(df['inc']))
plt.xlabel("Combined adult literacy rate (%) ages >=15; 2005-2014", fontsize=14)
plt.ylabel('Gross national income per capita (log scale)', fontsize=15)
plt.title('World Bank: Gross National Income Per Capita vs\n Combined Adult Literacy Rate');
```

<!-- BEGIN QUESTION -->

This scatter plot looks better. The relationship is closer to linear.

We can think of the log-linear relationship between x and y as follows: a constant change in x corresponds to a percent (scaled) change in y.

Next, try squaring literacy and taking the log of income. Don't forget to title the plot and label axes! Does the plot look better? 


```python
plt.figure(figsize=(10,5))
...
```

<!-- END QUESTION -->

Choosing the best transformation for a relationship is often a balance between keeping the model simple and straightening the scatter plot.

## Part 2: Kernel Density Estimation

In this part of the lab, you will develop a deeper understanding of how kernel density estimation works.

### Overview

Kernel density estimation is used to estimate a probability density function (i.e., a density curve) from a set of data. Just like a histogram, a density function's total area must sum to 1.

KDE centrally revolves around this idea of a "kernel". A kernel is a function whose area sums to 1. The three steps involved in building a kernel density estimate are:
1. Placing a kernel at each observation,
2. Normalizing kernels so that the sum of their areas is 1, and
3. Summing all kernels together.

The end result is a function that takes in some value `x` and returns a density estimate at the point `x`.

When constructing a KDE, there are several choices to make regarding the kernel. Specifically, we need to choose the function we want to use as our kernel, as well as a bandwidth parameter, which tells us how wide or narrow each kernel should be. We will explore these ideas now.

Suppose we have 3 data points with values 2, 4, and 9. We can compute the (useless) histogram with a KDE, as shown below. 

To understand how KDEs are computed, we need to see the KDE outside the given range. The easiest way to do this is to use an old function called `distplot` [(documentation)](https://seaborn.pydata.org/archive/0.11/generated/seaborn.distplot.html#seaborn.distplot). During the Spring 2024 offering of this course, `distplot` is still a working function in `seaborn`, but it will be removed at a future date. If you get an error that says that `distplot` is not a valid function, sorry, you are too far in the future to do this lab exercise.

```python
# Run this cell to create KDE plot and histogram using distplot; no further action is needed.
data3pts = np.array([2, 4, 9])
sns.distplot(data3pts, kde=True);
```

Using a deprecated function will often result in a `UserWarning`. You may follow their suggestions to adapt your code. In this case, we can use `displot` [(documentation)](https://seaborn.pydata.org/archive/0.11/generated/seaborn.displot.html#seaborn.displot) with some additional parameters. These additional parameters are needed because the default values of `displot` are different from `distplot`; we can manually set them to be the same.

```python
# Run this cell to create KDE plot and histogram using displot; no further action is needed.
plt.figure(figsize=(5, 15))
sns.displot(data3pts, kde=True, stat="density", kde_kws={"cut":4}, bins=2, height=5, aspect=1.5);
```

One question you might be wondering is how the kernel density estimator decides how "wide" each point should be. It turns out this is a parameter you can set called `bw`, which stands for bandwidth. For example, the code below gives a bandwidth value of 0.5 to each data point. You'll see the resulting KDE is quite different. Try experimenting with different values of bandwidth and see what happens.

```python
# Run this cell to plot distplot with the specified bandwidth. Try setting bw to different values.
sns.distplot(data3pts, kde=True, kde_kws={"bw": 0.5});
```

The equivalent plot using `displot` is provided below.

```python
# Run this cell to plot displot with the specified bandwidth. 
plt.figure(figsize=(5, 15))
sns.displot(data3pts, kde=True, stat="density", kde_kws={"bw_method": 0.5, "cut":4}, bins=2, height=5, aspect=1.5);
```

<br>

---
### Question 2a

As mentioned above, the kernel density estimate is just the sum of a bunch of copies of the kernel, each centered on our data points. The default kernel used by the `displot` function (as well as `kdeplot`) is the Gaussian kernel, given by:

$$\Large
K_\alpha(x, z) = \frac{1}{\sqrt{2 \pi \alpha^2}} \exp\left(-\frac{(x - z)^2}{2  \alpha ^2} \right)
$$

We've implemented the Gaussian kernel for you in Python below. Here, `alpha` is the smoothing or bandwidth parameter $\alpha$ for the KDE, `z` is the center of the Gaussian (i.e., a data point or an array of data points), and `x` is an array of values of the variable whose distribution we are plotting. In other words, `z` represents the center point of our smooth KDE bell curve, while `x` represents the range of values over which we want to generate the KDE plot.

```python
def gaussian_kernel(alpha, x, z):
    """
    Compute the Gaussian density estimate for values in x.

    Args:
        alpha: the smoothing parameter to pass to the kernel.
        x: an array of values whose density will be calculated.
        z: center of Gaussian.

    Returns:
        The smoothed estimate at values of x.
    """    
    return 1.0/np.sqrt(2. * np.pi * alpha**2) * np.exp(-(x - z) ** 2 / (2.0 * alpha**2))
```

For example, we can plot the Gaussian kernel centered at 9 with $\alpha$ = 0.5 as below: 

```python
xs = np.linspace(-2, 12, 200)
alpha = 0.5
kde_curve = [gaussian_kernel(alpha, x, 9) for x in xs]
plt.plot(xs, kde_curve);
```

<!-- BEGIN QUESTION -->

In the cell below, plot the 3 kernel density functions corresponding to our 3 data points on the same axis. Use an `alpha` value of 0.5. Recall that our three data points are 2, 4, and 9. 

**Note:** Make sure to normalize your kernels! This means that the area under each of your kernels should be $\frac{1}{3}$ since there are three data points.

You don't have to use the following hints, but they might be helpful in simplifying your code.

**Hint:** The `gaussian_kernel` function can also take a `NumPy` array as an argument for `z`.

**Hint:** To plot multiple plots at once, you can use `plt.plot(xs, y)` with a two-dimensional array as `y`.


<!-- grade_id: cell-fb3e2a6d258ddbf8, points: 1, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

In the cell below, we see a plot that shows the sum of all three of the kernels above. The plot resembles the KDE shown when you called `distplot` function with bandwidth 0.5 earlier. The area under the final curve will be 1 since the area under each of the three normalized kernels is $\frac{1}{3}$.


<!-- grade_id: cell-39b7921225a4bbda, points: 1, grade: true, solution: true -->
```python
# Run this cell to plot the sum of the kernels; no further action is needed.
xs = np.linspace(-2, 12, 200)
alpha = 0.5
kde_curve = np.array([1/3 * gaussian_kernel(alpha, x, data3pts) for x in xs])
plt.plot(xs, np.sum(kde_curve, axis=1));
```

Recall that earlier we plotted the kernel density estimation for the logarithm of the income data, as shown again below.

```python
# Run this cell to plot KDE of log income; no further action is needed.
ax = sns.displot(data=df, x=np.log10(df['inc']), kind='kde', rug=True)
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016', fontsize=13)
plt.title('World Bank Gross National Income Per Capita');
```

In the cell below, a similar plot is shown using what was done above. Try out different values of alpha in {0.1, 0.2, 0.3, 0.4, 0.5}. You will see that when `alpha=0.2`, the graph matches the previous graph well, except that the `displot` function hides the KDE values outside the range of the available data. Recall that `alpha` represents the spread of each individual kernel curve, which in turn affects the smoothness of the final KDE plot. When would you expect a smoother plot - when alpha = 0.1 or 0.5?

<!-- grade_id: cell-a1cbcf4340aacbdd, points: 1, grade: true, solution: true -->
```python
# Run this cell to plot KDE of log income; try out different bandwidths!
xs = np.linspace(1, 6, 200)
alpha = 0.5
kde_curve = np.array([1/len(df['inc']) * gaussian_kernel(alpha, x, np.log10(df['inc'])) for x in xs])
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016', fontsize=13)
plt.ylabel("Density")
plt.title('World Bank Gross National Income Per Capita')
plt.plot(xs, np.sum(kde_curve, axis = 1));
```

<br>

---
### Question 2b

In your answers above, you hard-coded a lot of your work. In this problem, you'll build a more general kernel density estimator function.

Implement the KDE function, which computes:

$$\Large
f_\alpha(x) = \frac{1}{n} \sum_{i=1}^n K_\alpha(x, z_i)
$$

where each $z_i$ represents a single datapoint in the collected dataset, $\alpha$ is a parameter to control the smoothness, and $K_\alpha$ is the kernel density function passed as `kernel`. Your code should run no longer than a couple of seconds. 


<!-- grade_id: q3a-answer, points: 0, grade: false, solution: true -->
```python
def kde(kernel, alpha, x, data):
    """
    Compute the kernel density estimate for the single query point x.

    Args:
        kernel: a kernel function with 3 parameters: alpha, x, and data.
        alpha: the smoothing parameter to pass to the kernel.
        x: a single query point (in one dimension).
        data: a NumPy array of data points.

    Returns:
        The smoothed estimate at the query point x.
    """    
    ...
kde(gaussian_kernel, 1.0, 2.0, np.array([3.0, 4.0, 5.0, 7.0]))
```

```python
grader.check("q2b")
```

Assuming you implemented `kde` correctly, the code below should generate the `kde` of the log of the income data as before.

```python
# Run this cell to generate the kde of the log of the income data; no further action is needed.
df['trans_inc'] = np.log10(df['inc'])
xs = np.linspace(df['trans_inc'].min(), df['trans_inc'].max(), 1000)
curve = [kde(gaussian_kernel, alpha, x, df['trans_inc']) for x in xs]
plt.hist(df['trans_inc'], density=True, color='orange')
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016', fontsize = 13);
plt.title('World Bank Gross National Income Per Capita')
plt.plot(xs, curve, 'k-');
```

And the code below should show a 3 x 3 set of plots showing the output of the kde for different `alpha` values.

```python
# Run this cell to generate the kde of the log of the income data 
# with different alphas values, no further action is needed.
plt.figure(figsize=(15,15))
alphas = np.arange(0.2, 2.0, 0.2)
for i, alpha in enumerate(alphas):
    plt.subplot(3, 3, i+1)
    xs = np.linspace(df['trans_inc'].min(), df['trans_inc'].max(), 1000)
    curve = [kde(gaussian_kernel, alpha, x, df['trans_inc']) for x in xs]
    plt.hist(df['trans_inc'], density=True, color='orange')
    plt.plot(xs, curve, 'k-')
    plt.title(r"$\alpha = " + format(alpha, ".02") + "$")
plt.show()
```

<br>

---
### Question 2c


Let's take a look at another kernel, the Boxcar kernel. The function `boxcar_kernel` is defined below.

\begin{equation} \Large
K_{a}(x, z) =
    \begin{cases}
        \frac{1}{\alpha}, & \text{if } -\frac{\alpha}{2} \leq (x-z) \leq \frac{\alpha}{2}\\
        0, & \text{otherwise}
    \end{cases}
\end{equation}

```python
def boxcar_kernel(alpha, x, z):
    """
    Compute the boxcar density estimate for values in x.

    Args:
        alpha: the smoothing parameter to pass to the kernel.
        x: an array of values whose density will be calculated.
        z: center of boxcar function.

    Returns:
        The smoothed estimate at values of x.
    """    
    return (((x-z)>=-alpha/2)&((x-z)<=alpha/2))/alpha
```

Run the cell below to enable interactive plots.

```python
from ipywidgets import interact
```

Now, we can plot the Boxcar and Gaussian kernel functions to see what they look like.

```python
x = np.linspace(-10,10,1000)
def f(alpha):
    plt.plot(x, boxcar_kernel(alpha,x,0), label='Boxcar')
    plt.plot(x, gaussian_kernel(alpha,x,0), label='Gaussian')
    plt.legend(title='Kernel Function')
    plt.show()
interact(f, alpha=(1,10,0.1));
```

Using the interactive plot below, compare the two kernel techniques on the data `data3pts` or log income. Assign `demo = 1` to see `data3pts` or `demo = 2` to see log income.

**Note:** Generating the KDE plot is slow, so expect some latency after you move the slider.

```python
demo = 1 # ... # set this value to 1 or 2

if demo == 1:
    xs = np.linspace(data3pts.min()-3, data3pts.max()+3, 1000)
    def f(alpha_g, alpha_b):
        plt.hist(data3pts, density=True, color='orange')
        g_curve = [kde(gaussian_kernel, alpha_g, x, data3pts) for x in xs]
        plt.plot(xs, g_curve, 'k-', label='Gaussian')
        b_curve = [kde(boxcar_kernel, alpha_b, x, data3pts) for x in xs]
        plt.plot(xs, b_curve, 'r-', label='Boxcar')
        plt.legend(title='Kernel Function')
        plt.show()
    interact(f, alpha_g=(0.01,.5,0.01), alpha_b=(0.01,3,0.1));
else:
    xs = np.linspace(df['trans_inc'].min(), df['trans_inc'].max(), 1000)
    def f(alpha_g, alpha_b):
        plt.hist(df['trans_inc'], density=True, color='orange')
        g_curve = [kde(gaussian_kernel, alpha_g, x, df['trans_inc']) for x in xs]
        plt.plot(xs, g_curve, 'k-', label='Gaussian')
        b_curve = [kde(boxcar_kernel, alpha_b, x, df['trans_inc']) for x in xs]
        plt.plot(xs, b_curve, 'r-', label='Boxcar')
        plt.legend(title='Kernel Function')
        plt.show()
    interact(f, alpha_g=(0.01,.5,0.01), alpha_b=(0.01,3,0.1));
```

<!-- BEGIN QUESTION -->

Briefly compare and contrast the Gaussian and Boxcar kernels in the cell below. How do the two kernels relate with each other for the same alpha value?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Part 3: Seaborn's `sns.kdeplot` 
Now that you've explored how KDE is calculated and how it compares to other plots like histograms, scatterplots, boxcar kernels, and rugplots, it's time to practice using `sns.kdeplot` [(documentation)](https://seaborn.pydata.org/generated/seaborn.kdeplot.html)! This lab will explore the 1-dimensional plotting capabilities of `kdeplot` while Homework 4 will explore its 2-dimensional plotting capabilities. 

<!-- BEGIN QUESTION -->

---
### Question 3a
Let's revisit the **income per capita** data from `q1b`. In the cell below, use `sns.kdeplot` to create a KDE plot of the income per capita. Don't forget to title the plot and label axes!

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

--- 

### Question 3b
We can see that our curve extends into negative values of our x-axis, which doesn't fully make sense for a variable like the gross national income per capita. This happens because there are a lot of data points concentrated near 0, resulting in a plot that is skewed right. Transforming the plot logarithmically will give us a more symmetric distribution of values that can make it easier to see patterns.

Similar to `q1c`, make a KDE plot that is scaled logarithmically (base 10). However, instead of manually transforming the data with `np.log10` and assigning the result to the `x=` argument, set the `kdeplot` arguments `x='inc'` and `log_scale=True`.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

--- 

### Question 3c

Compare your KDE plot from `q3b` to the one from `q1c` shown below: 

<img src='q1c.png' width="600px" alt="plot from q1c"/>

Despite the graphs having different x-axes, the shape of their KDE plots look very similar. Why do you think that is? 

**Hint**: Notice the units and scale of the x-axes. Where is the log applied? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Bike Sharing

## 📜 Introduction

Bike-sharing systems are a new(ish!) generation of traditional bike rentals where the process of signing up, renting, and returning is automated. Using these systems, users can easily rent a bike from one location and return it to another. 

In this assignment, you will clean, visualize, explore, and analyze bike-sharing data from Washington, D.C. You will also investigate open-ended questions. These open-ended questions ask you to think critically about how the plots you have created provide insight into the data.

After completing this assignment, you should be comfortable with:

* Reading plain text delimited data into `pandas`,
* Wrangling data for analysis,
* Using EDA to learn about your data, and 
* Making informative plots.

```python
# Run this cell to set up your notebook. Make sure ds100_utils.py is in this assignment's folder
import csv
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import zipfile
from pathlib import Path
import ds100_utils

# Default plot configurations
%matplotlib inline
plt.rcParams['figure.figsize'] = (16,8)
plt.rcParams['figure.dpi'] = 150
sns.set()

import warnings
warnings.filterwarnings("ignore")

from IPython.display import display, Latex, Markdown
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 💽 Loading Bike Sharing Data
We will explore data collected from a bike-sharing system in Washington, DC.

Here are variables in the data:

Variable       | Description
-------------- | ------------------------------------------------------------------
instant | unique ID
dteday | date
season | 1. spring <br> 2. summer <br> 3. fall <br> 4. winter
yr | year (0: 2011, 1:2012)
mnth | month (1 to 12)
hr | hour (0 to 23)
holiday | whether the day is a holiday or not
weekday | day of the week
workingday | if the day is neither a weekend nor a holiday
weathersit | 1. clear or partly cloudy <br> 2. mist and clouds <br> 3. light snow or rain <br> 4. heavy rain or snow
temp | normalized temperature in Celsius (divided by 41)
atemp | normalized "feels-like" temperature in Celsius (divided by 50)
hum | normalized percent humidity (divided by 100)
windspeed| normalized wind speed (divided by 67)
casual | count of casual users
registered | count of registered users
cnt | count of total rental bikes, including casual and registered  

###  📥 Download the Data

```python
# Run this cell to download the data; no further action is needed.

data_url = 'https://github.com/DS-100/fa20/raw/gh-pages/resources/assets/datasets/hw3-bikeshare.zip'
file_name = 'data.zip'
data_dir = 'data/'

dest_path = ds100_utils.fetch_and_cache(data_url=data_url, data_dir=data_dir, file=file_name)
print('Saved at {}'.format(dest_path))

zipped_data = zipfile.ZipFile(dest_path, 'r')

data_dir = Path('data')
zipped_data.extractall(data_dir)


print("Extracted Files:")
for f in data_dir.glob("*"):
    print("\t",f)
```

### 🔎 Examining the File Contents

As a quick, ungraded exercise, try and see if you can identify the file format of the file below. (No answer required).

```python
# Run this cell to look at the top of the file; no further action is needed.
for line in ds100_utils.head(data_dir/'bikeshare.txt'):
    print(line,end="")
```

### 🗻 Size
Is the file big?  How many records do we expect to find? Once again, this is ungraded, and there is no answer required.

```python
# Run this cell to view some metadata; no further action is needed.
print("Size:", (data_dir/"bikeshare.txt").stat().st_size, "bytes")
print("Line Count:", ds100_utils.line_count(data_dir/"bikeshare.txt"), "lines")
```

### 💾 Loading the Data

The following code loads the data into the `bike` `DataFrame`.

```python
# Run this cell to load the data; no further action is needed.
bike = pd.read_csv(data_dir/'bikeshare.txt')
bike.head()
```

Below, we show the shape of the file. 

You should see that the size of the `DataFrame` matches the number of lines in the file minus the header row.

```python
bike.shape
```

<hr style="border: 1px solid #fdb515;" />

## Question 4: Examining the Data

Before we start working with the data, let's examine its granularity.

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 4a
Granularity refers to the level of detail in a dataset—what each row represents in terms of time, space, or entity. In this dataset, each row corresponds to **bike-sharing data per hour** in Washington, DC. Based on the granularity and the variables present in the data, what might be some of the limitations of using this data? 

What are two additional data categories/variables that one could collect to address some of these limitations?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Use this cell for scratch work. If you need to add more cells for scratch work, add them BELOW this cell.
```

<hr style="border: 1px solid #fdb515;" />

## Question 5: Data Preparation
Several variables in the dataset are numeric/integer but actually represent categorical data. These include `holiday`, `weekday`, `workingday`, and `weathersit`.

In this problem, we will convert these four variables to strings that specify the categories. Specifically, use 3-letter labels (`Sun`, `Mon`, `Tue`, `Wed`, `Thu`, `Fri`, and `Sat`) for `weekday`. For `holiday` and `workingday`, use `yes`/`no`.

We will *mutate* the `bike` `DataFrame`, which means we will overwrite the corresponding variables in `bike`. This is an in-place data transformation, so the changes will be documented in the notebook for future reference. Ensure that the underlying data file `bikeshare.txt` remains unmodified.


<br>

--- 
### Question 5a 


Decode the `holiday`, `weekday`, `workingday`, and `weathersit` fields:

1. `holiday`: Convert to `yes` and `no`.  **Hint:** There are fewer holidays.
1. `weekday`: Mutate the `'weekday'` column to use the 3-letter label (`'Sun'`, `'Mon'`, `'Tue'`, `'Wed'`, `'Thu'`, `'Fri'`, and `'Sat'`) instead of its current numerical values. Assume `0` corresponds to `Sun`, `1` to `Mon`, and so on, in order of the previous sentence.
1. `workingday`: Convert to `yes` and `no`.
1. `weathersit`: You should replace each value with one of `Clear`, `Mist`, `Light`, or `Heavy`. Assume `1` corresponds to `Clear`, `2` corresponds to `Mist`, and so on in order of the previous sentence.

**Note:** If you mutate any of the fields above, then they will not be in the format of their original `.txt` file. As a debugging tip, if you want to revert changes, rerun the cell above that loads in the `bike` `DataFrame`.

**Hint:** One approach is to use the `replace`[(documentation)](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.replace.html) method of the `pd.DataFrame` class.

<!-- grade_id: q1a-answer, points: 0, grade: false, solution: true -->
```python
# Modify holiday, weekday, workingday, and weathersit here.
...
bike = ...
bike.head()
```

```python
grader.check("q5a")
```

<br>

--- 
### Question 5b
In the next few questions, we will analyze the daily number of registered and unregistered users.

Construct a `DataFrame` named `daily_counts` indexed by `dteday` with the following columns:
* `casual`: Total number of casual riders for each day.
* `registered`: Total number of registered riders for each day.
* `workingday`: Whether that day is a working day or not (`yes` or `no`).



<!-- grade_id: q1c-answer, points: 0, grade: false, solution: true -->
```python
daily_counts = ...
daily_counts.head()
```

```python
grader.check("q5b")
```

<hr style="border: 1px solid #fdb515;" />

## Question 6: Exploring the Distribution of Riders

Let's begin by comparing the distribution of the daily counts of casual and registered riders. Questions 6-10 require using many visualization methods, so for your convenience, we have summarized a few useful ones below.

### Matplotlib and Seaborn Table of Common Functions

`x` and `y` are sequences of values (i.e., arrays, lists, or `Series`).

| Function | Description |
| -------- | ----------- |
| `plt.title(name)` | Adds a title `name` to the current plot. |
| `plt.xlabel(name)` | Adds a label `name` to the x-axis. |
| `plt.ylabel(name)` | Adds a label `name` to the y-axis. |
| `plt.plot(x, y)` | Creates a line plot of `x` against `y`. |
| `plt.scatter(x, y)` | Creates a scatter plot of `x` against `y`. |
| `plt.hist(x, bins=None)` | Creates a histogram of `x`; `bins` can be an integer or a sequence. |
| `plt.bar(x, height)` | Creates a bar plot of categories `x` and corresponding heights `height`. |
| `sns.histplot(data=, x=, y=, hue=, kde=)` | Creates a distribution plot; `data` is a `DataFrame`; `x`, `y` are column names in `data` that specify positions on the `x` and `y` axes; `hue` is a column name in `data` that adds subcategories to the plot based on `hue`; `kde` is a boolean that determines whether to overlay a KDE curve. |
| `sns.lineplot(data=, x=, y=, hue=)` | Creates a line plot. |
| `sns.scatterplot(data=, x=, y=, hue=, size=)` | Creates a scatter plot; `size` is a vector that contains the size of the point for each subcategory based on `hue`. |
| `sns.kdeplot(data=, x=, y=)` | Creates a kernel density estimate plot; `data` is a `DataFrame`; `x`, `y` are column names in `data` that specify positions on the `x` and `y` axes. |
| `sns.jointplot(data=, x=, y=, kind=)` | Creates a joint plot of 2 variables with a KDE plot in the middle and a distribution plot for each variable on the sides; `kind` determines the visualization type for the distribution plot, can be `scatter`, `kde` or `hist`. |

**Note**: This list of functions and parameters is **not** exhaustive. You may need to reference and explore more documentation to answer the following questions, but we will help you through that process.

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 6a

Use the `sns.histplot`[(documentation)](https://seaborn.pydata.org/generated/seaborn.histplot.html#seaborn.histplot) function to create a plot that overlays the distribution of the daily counts of bike users.

- Use blue to represent `casual` riders, and red to represent `registered` riders. 

The temporal granularity of the records should be daily counts, which you should have after completing question 5c. In other words, you should be using `daily_counts` to answer this question.

**Hints:** 
- You will need to set the `stat` parameter appropriately to match the desired plot.
- The `label` parameter of `sns.histplot` allows you to specify, as a string, how the plot should be labeled in the legend. Although label is not explicitly documented in Seaborn, it works because `sns.histplot` internally relies on  `matplotlib`, which supports the label parameter. For example, passing in `label="My data"` would give your plot the label `"My data"` in the legend.
- You will need to make two calls to `sns.histplot`.
 

Include a `legend`, `xlabel`, `ylabel`, and `title`. Read the [seaborn plotting tutorial](https://seaborn.pydata.org/tutorial/distributions.html) if you're not sure how to add these. After creating the plot, look at it and make sure you understand what the plot is actually telling us, e.g., on a given day, the most likely number of registered riders we expect is ~4000, but it could be anywhere from nearly 0 to 7000.

<img src='images/casual_v_registered.png' width="700px" alt="Visualization showing the difference between casual and registered bikers"/>

For all visualizations in DSC 100, we will evaluate your plot based on its similarity to the provided example. While your plot does not need to be _identical_ to the example shown, we do expect it to capture its main features, such as the **general shape of the distribution**, the **axis labels**, the **legend**, and the **title**. It is okay if your plot contains small stylistic differences, such as differences in color, line weight, font, or size/scale.

<!-- grade_id: q2a-answer, points: 2, grade: true, solution: true -->
```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 6b

In the cell below, describe the differences you notice between the density curves for casual and registered riders.  

- Consider concepts such as modes, symmetry, skewness, tails, gaps, and outliers. 

- Include a comment on the spread of the distributions. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 6c

The density plots do not show us how the counts for `registered` and `casual` riders vary together. 

Use `sns.lmplot` [(documentation)](https://seaborn.pydata.org/generated/seaborn.lmplot.html) to create a scatter plot to investigate the relationship between casual and registered counts. 

- Use the `bike` `DataFrame` to plot hourly counts instead of daily counts.
- Color the points in the scatter plot according to whether or not the day is a working day. Your colors do not have to match ours exactly, but they should be different based on whether the day is a working day. 

<img src='images/casual_registered_working_nonworking.png' width="800px"  alt="Scatterplot showing the difference between casual and registered bikers"/>

**Hints:** 
* Check out this helpful [tutorial on `lmplot`](https://seaborn.pydata.org/tutorial/regression.html).
* There are many points in the scatter plot, so make them small to help reduce overplotting. Check out the `scatter_kws` parameter of `lmplot`.
* You can set the `height` parameter if you want to adjust the size of the `lmplot`.
* Add a descriptive title and axis labels for your plot.
* It is okay if the scales of your `x` and `y` axis (i.e., the numbers labeled on the two axes) are different from those used in the provided example.


<!-- grade_id: q2c-answer, points: 2, grade: true, solution: true -->
```python
sns.set(font_scale=1) # This line automatically makes the font size a bit bigger on the plot. You do not need to change this line
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 6d

What does this scatterplot seem to reveal about the relationship (if any) between casual and registered riders and whether or not the day is on the weekend? 

What effect does overplotting have on your ability to describe this relationship?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 7: Visualization
 
To address overplotting, let's try visualizing the data with another technique, the bivariate kernel density estimate. You will want to read up on the documentation for `sns.kdeplot`, which can be found [here](https://seaborn.pydata.org/generated/seaborn.kdeplot.html). 

The result we wish to achieve should be a plot that looks like this:

<img src='images/bivariate_kde_of_daily_rider_types_simple.png' width="600px" alt="Kernel density plot showing the difference between riders"/>

A basic KDE plot of all the data is quite easy to generate (see below). However, this plot combines both weekend and weekday data, which isn't what we want (see the example figure above).

```python
sns.kdeplot(data=daily_counts, x='casual', y='registered')
plt.title('KDE Plot Comparison of Registered vs Casual Riders');
```

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 7a

Generate a bivariate kernel density plot with workday and non-workday separated using the `daily_counts` `DataFrame`. It should look like the first plot displayed above.

**Hint:** You only need to call `sns.kdeplot` once. Take a look at the `hue` parameter and adjust other inputs as needed.

After you get your plot working, experiment by setting `fill=True` in `kdeplot` to see the difference between the shaded and unshaded versions.

- But, **please submit your work with `fill=False`.**

```python
# Set the figure size for the plot
plt.figure(figsize=(12,8))

...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 7b
<img src='images/bivariate_kde_of_daily_rider_types_gradient.png' width="800px" alt="fancy KDE plot"/>

With some modification to your Question 7a code (this modification is not in scope), we can generate the plot above. 

In your own words, describe what the lines and the color shades of the lines signify about the data. What does each line and color represent? 

**Hint**: You may find it helpful to compare it to a contour or topographical map as shown [here](https://gisgeography.com/contour-lines-topographic-map/).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 7c

What additional details about the riders can you identify from this contour plot that were difficult to determine from the scatter plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 8: Understanding Daily and Monthly Patterns

<br>

--- 

### Question 8a
Create a new column in the `bike` `DataFrame` called `time_category`, which maps the `hr` column into the following broader time categories:

- **Midnight**: 2 AM to before 5 AM  
- **Morning**: 5 AM to before 11 AM  
- **Lunch Time**: 11 AM to before 2 PM  
- **Afternoon**: 2 PM to before 5 PM  
- **Evening**: 5 PM to before 9 PM  
- **Night**: 9 PM to before 2 AM  

Then, display the first few rows of the updated `DataFrame` to verify that the new column is correctly added.

```python
...

bike["time_category"] = ...

bike.head()
```

```python
grader.check("q8a")
```

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 8b

Let’s examine the behavior of riders by plotting the **average number of riders** for each **time category** (using the `time_category` column), separated by rider type.

Your plot should look like the plot below. It's fine if your plot's colors don't match ours exactly.

**Hint:**  
To label the x-axis correctly, use `plt.xticks()` to manually set tick positions and labels. You may need to rotate the labels for readability. Refer to the [documentation](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.xticks.html) for more details.

<img src='images/time_category_bikes.png' width="700px" alt="Riders according to time"/>                                                                   

```python
# Group by time category and calculate means
time_category_means = (
    bike.groupby("time_category")[["casual", "registered"]].mean()
)

plt.figure(figsize=(10, 7))
sns.lineplot(
    data=time_category_means.reset_index(),
    x="time_category",
    y="casual",
    label="Casual Riders",
)
sns.lineplot(
    data=time_category_means.reset_index(),
    x="time_category",
    y="registered",
    label="Registered Riders",
)

plt.xlabel("Time of the Day (Categories)")
plt.ylabel("Average Count")
plt.title("Average Count of Casual vs. Registered Riders by Time Categories")
plt.xticks(
    ticks=range(len(time_category_means)),  # Order categories
    labels=["Midnight", "Morning", "Lunch Time", "Afternoon", "Evening", "Night"],
    rotation=45 # Rotate x-axis labels for readability
)  
plt.legend()
plt.tight_layout()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 8c

Next, analyze how the average count of casual and registered riders varies by month (`mnth`). 

Compute the average number of casual and registered riders for each month in the dataset and create a line plot showing the trends.

Your plot should look like the plot below. It's fine if your plot's colors don't match ours exactly.


<img src='images/monthly_avg_riders.png' width="700px" alt="Plot showing the difference between causal and registered riders by month"/>


```python
# Group by month and calculate mean rider counts
avg_riders_by_month = bike.groupby("mnth")[["casual", "registered"]].mean()

plt.figure(figsize=(10, 7))

# Plot casual riders
sns.lineplot(
    data=avg_riders_by_month.reset_index(),
    x="mnth",
    y="casual",
    label="Casual Riders"
)

# Plot registered riders
sns.lineplot(
    data=avg_riders_by_month.reset_index(),
    x="mnth",
    y="registered",
    label="Registered Riders"
)

# Formatting
plt.xlabel("Month")
plt.ylabel("Average Rider Count")
plt.title("Average Number of Casual vs. Registered Riders by Month")
plt.xticks(
    ticks=range(1, 13),  # Months range from 1 to 12
    labels=[
        "Jan", "Feb", "Mar", "Apr", "May", "Jun",
        "Jul", "Aug", "Sep", "Oct", "Nov", "Dec"
    ],
    rotation=45  # Rotate x-axis labels for readability
)
plt.legend()
plt.tight_layout()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 8d

What can you observe from the plots generated in **8b** and **8c**? 

Discuss your observations for both types of riders, and hypothesize about the meaning of the peaks and troughs of both riders' distributions.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 9: Exploring Ride-Sharing and Weather
Now let's examine how the weather may be associated. First, let's look at how the proportion of casual riders changes as the weather changes.

<br>

--- 

### Question 9a
Create a new column `prop_casual` in the `bike` `DataFrame` representing the proportion of casual riders out of all riders for each record.

<!-- grade_id: q6a-answer, points: 0, grade: false, solution: true -->
```python
...
```

```python
grader.check("q9a")
```


In order to examine the relationship between the proportion of casual riders and temperature, we can create a scatterplot using [`sns.scatterplot`](https://seaborn.pydata.org/generated/seaborn.scatterplot.html). 

- We can even use color/hue to encode the information about the day of the week. 

Run the cell below, and you'll see we end up with a big mess that is impossible to interpret.

```python
plt.figure(figsize=(10, 7))
sns.scatterplot(data=bike, x='temp', y='prop_casual', hue='weekday');
```

We could attempt linear regression using `sns.lmplot` as shown below, which hints at some relationships between temperature and proportional casual, but the plot is still fairly unconvincing. 

```python
sns.lmplot(data=bike, x='temp', y='prop_casual', hue='weekday', scatter_kws={"s": 20}, height=10)
plt.title("Proportion of Casual Riders by Weekday");
```

A better approach to visualize the relationship between temperature and the proportion of casual riders is to use local smoothing techniques.

- The basic idea is to compute a representative `y` value for each `x` value that captures the data close to that `x` value.

- One such technique is "Locally Weighted Scatterplot Smoothing" (LOWESS/LOESS). 

Below is an example where the red curve represents a smoothed pattern in the scatterplot.

```python
from statsmodels.nonparametric.smoothers_lowess import lowess
# Make noisy data
x_obs = np.sort(np.random.rand(100)*4.0 - 2)
y_obs = np.exp(x_obs) + np.random.randn(100) / 2.0

# Plot a regular scatterplot with our noisy data 
sns.scatterplot(x=x_obs, y=y_obs, label="Raw Data")

# Pass in the noisy data to lowess to generate the smoothed y-values, 
# we want to preserve the relative ordering of the y-values, which is why return_sorted=False
y_smooth = lowess(y_obs, x_obs, return_sorted=False)

# Plot a red line using the smoothed y-values
sns.lineplot(x=x_obs, y=y_smooth, label="Smoothed Estimator", color='red')
plt.legend();
```

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 9b

Draw 7 smoothed curves on a single plot, one for each day of the week. 

- The x-axis should be the temperature (as given in the `'temp'` column).

- The y-axis should be a smoothed version of the proportion of casual riders. 

<img src="images/curveplot_temp_prop_casual.png" width="600px" alt="Temperature vs Casual Rider Proportion by Weekday"/>

You should use [`statsmodels.nonparametric.smoothers_lowess.lowess`](http://www.statsmodels.org/dev/generated/statsmodels.nonparametric.smoothers_lowess.lowess.html) just like the example above. 

- Unlike the example above, plot ONLY the lowess curve. Do not plot the actual data, which would result in overplotting. 

You do not need to match the colors on our sample plot as long as the colors in your plot make it easy to distinguish which day they represent.

**Hints:** 
* Start by plotting only one day of the week to make sure you can do that first. Then, consider using a `for` loop to repeat this plotting operation for all days of the week.

* The `lowess` function expects the `y` coordinate first, then the `x` coordinate. You should also set the `return_sorted` field to `False`.

* **You will need to rescale the normalized temperatures stored in this dataset to Fahrenheit values.** Look at the section of this notebook titled 'Loading Bike Sharing Data' for a description of the (normalized) temperature field to know how to convert back to Celsius first. After doing so, convert it to Fahrenheit. By default, the temperature field ranges from 0.0 to 1.0. In case you need it, $\text{Fahrenheit} = \text{Celsius} \times \frac{9}{5} + 32$. If you prefer plotting temperatures in Celsius, that's fine as well! Just remember to convert accordingly so the graph is still interpretable. In addition, for smoother curves, use `sns.lineplot` instead of Matplotlib's default plotting functions.  
This helps avoid "noisy" jagged lines that might appear with `plt.plot` or `plt.scatter`.

<!-- grade_id: cell-831d12249d05ea2c, points: 4, grade: true, solution: true -->
```python
from statsmodels.nonparametric.smoothers_lowess import lowess

plt.figure(figsize=(10,8))
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 9c
Examine the plot above and describe how casual ridership changes with temperature. Determine if the **plot alone** provides evidence of a **causal** relationship between temperature and casual ridership, and explain your reasoning. 

Finally, based on **your own intuition**, state whether you think there is a underlying causal relationship. Justify your answer. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 10: Expanding our Analysis

<!-- BEGIN QUESTION -->

<br>

--- 

### Question 10a

Imagine you are working for a bike-sharing company that collaborates with city planners, transportation agencies, and policymakers in order to implement bike-sharing in a city. These stakeholders would like to reduce congestion and lower transportation costs. They also want to ensure the bike-sharing program is implemented equitably. In this sense, equity is a social value that informs the deployment and assessment of your bike-sharing technology. 

Equity in transportation includes: Improving the ability of people of different socio-economic classes, genders, races, and neighborhoods to access and afford transportation services and assessing how inclusive transportation systems are over time. 

Do you think the `bike` data as it is can help you assess equity? If so, please explain. If not, how would you change the dataset? 

You may discuss how you would change the granularity, what other kinds of variables you'd introduce to it, or anything else that might help you answer this question.

**Note**: There is no single "right" answer to this question – we are looking for thoughtful reflection and commentary on whether or not this dataset, in its current form, encodes information about equity.


_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Use this cell for scratch work. If you need to add more cells for scratch work, add them BELOW this cell.
```

<!-- BEGIN QUESTION -->

<br>

--- 
### Question 10b
[Bike sharing is growing in popularity](https://www.bts.gov/newsroom/bike-share-stations-us), and new cities and regions are making efforts to implement bike-sharing systems that complement their other transportation offerings. The [goals of these efforts](https://www.wired.com/story/americans-falling-in-love-bike-share/) are to have bike sharing serve as an alternate form of transportation in order to alleviate congestion, provide geographic connectivity, reduce carbon emissions, and promote inclusion among communities.

Bike-sharing systems have spread to many cities across the country. The company you work for asks you to determine the feasibility of expanding bike sharing to additional cities in the US. 

Based on your plots in this assignment, would you recommend expanding bike sharing to additional cities in the US? If so, what cities (or types of cities) would you suggest? 

Please list at least two reasons why, and mention which plot(s) you drew your analysis from. 

**Note**: There isn't a set right or wrong answer for this question. Feel free to come up with your own conclusions based on evidence from your plots! 


_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Use this cell for scratch work. If you need to add more cells for scratch work, add them BELOW this cell.
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Lab 6!


### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Submit this zip file to the Lab 6 Coding assignment on Gradescope. Unzip the file to obtain the pdf with your written answers and submit that pdf to the Lab 6 Writing Assignment on Catcourses.

**Important**: Please check that your **plots/graphs and written responses** were generated and submitted correctly to the Lab 6 Written Assignment.

**You are responsible for ensuring your submission follows our requirements and that the PDF for Lab 6 written answers was generated/submitted correctly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out prior to the deadline. 

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2b (points: 1)
```
>>> np.isclose(kde(gaussian_kernel, 1.0, 2.0, np.array([3.0, 4.0, 5.0, 7.0])), 0.075099)
True
```
```
>>> np.isclose(kde(gaussian_kernel, 2, 5, np.array([1.0, 3.0, 4.0, 5.0, 7.0, 100])), 0.107411)
True
```
```
>>> np.isclose(kde(gaussian_kernel, 100, 370, df['inc']), 0.000248205)
True
```


Tests for q5a (points: 3)
```
>>> isinstance(bike, pd.DataFrame)
True
```
```
>>> bike['holiday'].dtype == np.dtype('O')
True
```
```
>>> list(bike['holiday'].iloc[370:375]) == ['no', 'no', 'yes', 'yes', 'yes']
True
```
```
>>> bike['weekday'].dtype == np.dtype('O')
True
```
```
>>> bike['workingday'].dtype == np.dtype('O')
True
```
```
>>> bike['weathersit'].dtype == np.dtype('O')
True
```


Tests for q5b (points: 3)
```
>>> bool(np.all(daily_counts.sort_index().index[0:5] == ['2011-01-01', '2011-01-02', '2011-01-03', '2011-01-04', '2011-01-05']))
True
```
```
>>> bool(type(daily_counts) == pd.DataFrame)
True
```
```
>>> bool(np.isclose(np.round(daily_counts['casual'].var()), 471450))
True
```


Tests for q8a (points: 1)
```
>>> 'time_category' in bike.columns
True
```
```
>>> expected_categories = {'Midnight', 'Morning', 'Lunch Time', 'Afternoon', 'Evening', 'Night'}
>>> set(bike['time_category'].unique()) == expected_categories
True
```
```
>>> len(bike) == len(bike.drop(columns=['time_category']))
True
```
```
>>> expected_categories = {'Midnight', 'Morning', 'Lunch Time', 'Afternoon', 'Evening', 'Night'}
>>> bike['hr'].apply(hour_to_category).isin(expected_categories).all()
True
```


Tests for q9a (points: 1)
```
>>> int(bike['prop_casual'].sum()) == 2991
True
```
